---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example **exercise**.

`@instructions`
Make sure python is installed

`@hint`
Use the plus symbol to add add two numbers

`@pre_exercise_code`
```{python}
a = 2
b = 3
sum = a + b
print("sum: ", sum)
```

`@sample_code`
```{python}

```

`@solution`
```{python}
You should see sum: 5 printed in the console
```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 05ed3c8a9f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
